Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Orion"
COLLABORATORS = ""
```

---

# PA1

Our first programming assignment!

Covers intro to python: strings, if, while, range, for and lists

Due Friday, 9/20 by 5pm. 


# Problem 0 (10 points)

Given a string, return its length

Hints: <br>
    https://docs.python.org/3/library/stdtypes.html?highlight=replace#str.replace <br>
    https://docs.python.org/3/tutorial/introduction.html#strings

<!-- grade_id: strfn0-fn, points: 0, grade: false, solution: true -->
```python
def strfn0(s):
    """Given a string, return its length
    """
    # YOUR CODE HERE
    if len(s) <= 0:
        #raise IndexError
        return print(f"STRING IS EMPTY YOU DINGUS")
    else:
        print(f"The len of the string is {len(s)}")
        return len(s)
        
    raise NotImplementedError()
```

Expected output for strfn0("doodle") is 6

Use the cell below to test your function.

```python
#?str.replace()
strfn0("doodle")
```

Output:
```
The len of the string is 6
6
```

<!-- grade_id: strfn0-test0, points: 6, grade: true, solution: false -->
```python
"""Check that method returns the correct output for several inputs"""
assert strfn0("doodle") == 6
assert strfn0("Supercalifragilisticexpialidocious") == 34
```

Output:
```
The len of the string is 6
The len of the string is 34
```

<!-- grade_id: strfn0-validation, points: 4, grade: true, solution: false -->
```python
"""Check that function raises an error for invalid inputs"""
try:
    strfn0("")
except IndexError:
    raise AssertionError("did not handle empty string")
else:
    pass
```

Output:
```
STRING IS EMPTY YOU DINGUS
```

## Problem 1 (15 points)

Given a string, replace all the instances of the first char with *

Constraint: Do NOT use any for-loops in your solution. Employ string indexing or str methods only.

Keep everything in lower case.

Hint: See
    https://docs.python.org/3/library/stdtypes.html?highlight=replace#str.replace
    https://docs.python.org/3/tutorial/introduction.html#strings

<!-- grade_id: strfn1-fn, points: 0, grade: false, solution: true -->
```python
def strfn1(s):
    """Given a string, replace all the instances of the first char with *
    Constraint: No for-loops
    """
    # YOUR CODE HERE
    if len(s) <= 0:
        #raise IndexError
        return print(f"THE STRING IS EMPTY YOU DINGUS!")
    lowstr = s.lower()
    
    nwstr = lowstr[1:].replace(lowstr[0],"*",)
    
    return lowstr[0]+nwstr
    
    raise NotImplementedError()
```

Expected output for strfn1("Doodle") is "doo*le"

Use the cell below to test your function.

```python
#?str.replace()
strfn1("Supercalifragilisticexpialidocious")
```

Output:
```
'supercalifragili*ticexpialidociou*'
```

<!-- grade_id: strfn1-test0, points: 10, grade: true, solution: false -->
```python
"""Check that method returns the correct output for several inputs"""
assert strfn1("Doodle") == "doo*le"
assert strfn1("Supercalifragilisticexpialidocious") == "supercalifragili*ticexpialidociou*"
```

<!-- grade_id: strfn1-validation, points: 5, grade: true, solution: false -->
```python
"""Check that function raises an error for invalid inputs"""
try:
    strfn1("")
except IndexError:
    raise AssertionError("did not handle empty string")
else:
    pass
```

Output:
```
THE STRING IS EMPTY YOU DINGUS!
```

```python
# ^^ Ask if this is correct validation in class
```

## Problem 2 (20 points)

Given two digits: x and n, return $x+xx+xxx+...+x^n$

<!-- grade_id: strfn2-fn, points: 0, grade: false, solution: true -->
```python
def strfn2(x, n):
    """Given two digits: x and n, return $x+xx+xxx+...+x^n$
    """
    # YOUR CODE HERE
    
    # cast str to (int) for math
    #str to contactination
    xsum = 0
    while n != 0:
        xstr = str(x)*n
        xsum += int(xstr)
        n -= 1
    return xsum

    raise NotImplementedError()
```

Expected output for strfn2(2, 3) is 2 + 22 + 222 = 246

Use the cell below to test your function.

```python
#?str.replace()
strfn2(0,0)
```

Output:
```
0
```

<!-- grade_id: strfn2-test0, points: 20, grade: true, solution: false -->
```python
"""Check that method returns the correct output for several inputs"""
assert strfn2(2,3) == 246
assert strfn2(1,9) == 123456789
strfn2(3,3) == 369
```

Output:
```
True
```

## Problem 3 (25 points)

Write a function that
- takes 4 inputs: a,b, r0, r1
- finds all digits that are divisble by `a` BUT not by `b` and are in range `[r0, r1]`

<!-- grade_id: divisible-fn, points: 0, grade: false, solution: true -->
```python
def divisible(a, b, r0, r1):
    """ Write a program that takes 4 inputs: a,b, r0, r1.
    finds all that are divisble by a BUT not by b and are in range [r0, r1]
    """
    # YOUR CODE HERE
    if a <= 0 or b <= 0:
        raise ValueError
    else:
        inarr = []
        for num in range(r0,r1):
            inarr.append(num+1)
        outarr = []
        for num in inarr:
            if num%a==0 and num%b!=0:
                outarr.append(num)
        return outarr
```

Expected output for a=3, b=4, r0=0, r1=15 is [3,6,9,15]

Use the cell below to test your function.

```python
divisible(3, 4, 0, 15)

```

Output:
```
[3, 6, 9, 15]
```

<!-- grade_id: divisible-test0, points: 20, grade: true, solution: false -->
```python
"""Check that method returns the correct output for several inputs"""
assert divisible(3, 4, 0, 15) == [3, 6, 9, 15]
```

<!-- grade_id: divisible-validation, points: 5, grade: true, solution: false -->
```python
"""Check that function raises an error for invalid inputs"""
try:
    divisible(0, 4, 0, 8)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    divisible(4, 0, 0, 8)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

## Problem 4 (30 points)

Write a function that
- takes 2 inputs: a,b
- use a nested list to create a 2D matrix that has `a` many rows and `b` many columns
- where each element equals to the sum of its indeces: matrix(i,j) = i+j


<!-- grade_id: matrixcreate-fn, points: 0, grade: false, solution: true -->
```python

def matrixcreate(a,b):
    # YOUR CODE HERE
    #matrix = [[a+b for num in range(a-1)] for y in range(b)]
    matrix = [[0]*b for num in range(a)]
    
    for i in range(a):
        for j in range(b):
            matrix[i][j] = i+j
    
    return matrix
    #raise NotImplementedError()
```

Hint:https://docs.python.org/3/tutorial/datastructures.html#nested-list-comprehensions


Expected output for a=3, b=2 is [[0, 1], [1, 2], [2, 3]]

Use the cell below to test your function.

```python
matrixcreate(3,2)
```

Output:
```
[[0, 1], [1, 2], [2, 3]]
```

<!-- grade_id: matrixcreate-test0, points: 30, grade: true, solution: false -->
```python
"""Check that method returns the correct output for several inputs"""
assert matrixcreate(3, 2) == [[0, 1], [1, 2], [2, 3]]
assert matrixcreate(2, 3) == [[0, 1, 2], [1, 2, 3]]

```

```python

```